---
title: 'Analyzing data'
description: 'Data frames contain data, arranged in columns.  In this chapter we get the column data and analyze it, using summary statistics and plots.'
---

## An introduction to columns

```yaml
type: NormalExercise
key: 12aa3a3b6f
xp: 100
```

Let us remind ourselves about the data.  Run the code in the code window.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}
gender_stats <- read.csv('http://bit.ly/gender-stats-data')
```

`@sample_code`
```{r}
head(gender_stats)
```

`@solution`
```{r}
head(gender_stats)
```

`@sct`
```{r}

```

## Another introduction to columns

```yaml
type: NormalExercise
key: 12aa3a3b6f
xp: 100
```

Let us remind ourselves again about the data.  Run the code in the code
window.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}
gender_stats <- read.csv('http://bit.ly/gender-stats-data')
```

`@sample_code`
```{r}
head(gender_stats)
```

`@solution`
```{r}
head(gender_stats)
```

`@sct`
```{r}

```
