---
title: 'Analyzing data'
description: 'Data frames contain data, arranged in columns.  In this chapter we get the column data and analyze it, using summary statistics and plots.'
---

## An introduction to columns

```yaml
type: NormalExercise
key: 12aa3a3b6f
xp: 100
```

Let us remind ourselves about the data.  Run the code in the code window.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Course Trailer for Causal Inference with R - Experiments

```yaml
type: VideoExercise
lang: r
xp: 50
skills: 1
video_link: //player.vimeo.com/video/280255549
```
