---
title: 'Working with columns'
description: 'Data frames contain data, arranged in columns.  In this chapter we get the column data and analyze it, using descriptive statistics.'
---

## Getting columns with dollars

```yaml
type: NormalExercise
key: 12aa3a3b6f
xp: 100
```

Let us remind ourselves about the data.  Run the code in the code window.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Columns as variables

```yaml
type: NormalExercise
key: 6bdb5ddcb1
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## What's missing in the mean?

```yaml
type: NormalExercise
key: a6dfce900b
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## More arguments

```yaml
type: NormalExercise
key: 0d8f1f12b2
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Calculating with columns

```yaml
type: NormalExercise
key: 08b638fb44
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## All-at-once calculations with columns

```yaml
type: NormalExercise
key: a811986c3c
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
