---
title: 'Analyzing data'
description: 'Data frames contain data, arranged in columns.  In this chapter we get the column data and analyze it, using summary statistics and plots.'
---

## An introduction to columns

```yaml
type: NormalExercise
key: 12aa3a3b6f
xp: 100
```

Let us remind ourselves about the data.  Run the code in the code window.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 2c845540dc
xp: 50
video_link: //player.vimeo.com/video/160339466
video_hls: //videos.datacamp.com/transcoded/672_intermediate_r/v2/hls-ch1_1.master.m3u8
```

`@projector_key`
fb1ff66627aae8d3992086b69b2746f9
