---
title: 'The R way'
description: 'R is a language that allows you to define your data analysis in words.   It is a completely different to something like Excel, where you do much of your analysis by clicking on the graphical user interface. '
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

R does not work the same way as packages like Excel.

In Excel, you use the mouse and the menus to perform your analysis.

In R, we use words to describe our analysis.

The language we use is not quite English, but it is fairly close, as you will see.

The language we are using is of course *R*.  R is a programming language, that is designed specifically for data analysis and statistics.

In this exercise, we do a very simple analysis.  We load some data, and have a look at the first few rows.  Next we plot one column against another, to look for a relationship.

`@instructions`
Click on Run code to run the analysis.  You will see the output in the Window on the right.

`@hint`


`@pre_exercise_code`
```{r}
data = read.csv('')
```

`@sample_code`
```{r}

```

`@solution`
```{r}
2 + 2
```

`@sct`
```{r}

```
