---
title: 'Looking at data with plots'
description: 'R has very powerful tools for looking at data with plots.  We explore some simple plots.'
---

## Histograms

```yaml
type: NormalExercise
key: 37d36434b1
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Histograms, variables

```yaml
type: NormalExercise
key: 32d8888562
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Scatter plots

```yaml
type: NormalExercise
key: 002cef1965
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Scatter plots with variables

```yaml
type: NormalExercise
key: f1fad864ed
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
